<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/data/Data_Analysis.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Kata: Data Analysis

Inspired by an activity in the [Open Machine Learning Course](https://mlcourse.ai/).

## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## Part 1 : Loading data

```python
# Import required libraries
import numpy as np
import pandas as pd
```

### Question

Use Pandas to load the dataset located [here](https://raw.githubusercontent.com/bpesquet/machine-learning-handbook/master/datasets/athlete_events.csv), which was extracted from this [Kaggle competition](https://www.kaggle.com/heesoo37/120-years-of-olympic-history-athletes-and-results/). Store it into a DataFrame named `df_olympics`.

<!-- grade_id: cell-6930168ca952d16e, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('https://raw.githubusercontent.com/bpesquet/machine-learning-handbook/master/datasets/athlete_events.csv', header=None).head()
df_olympics = pd.DataFrame(data)
```

<!-- grade_id: cell-a4b2d9d694a4d11b, points: 1, grade: true, solution: false -->
```python
print(f'df_olympics: {df_olympics.shape}')
assert isinstance(df_olympics, pd.DataFrame)
assert df_olympics.shape == (271116, 15)
```

## Part 2: Exploring data

### Question

Print general information about the DataFrame.

<!-- grade_id: cell-cc980d2fe4eb912e, points: 0, grade: false, solution: true -->
```python
print(df_olympics.info)
```

Output:
```
<bound method DataFrame.info of    0                     1    2    3       4       5               6    7   \
0  ID                  Name  Sex  Age  Height  Weight            Team  NOC   
1   1             A Dijiang    M   24     180      80           China  CHN   
2   2              A Lamusi    M   23     170      60           China  CHN   
3   3   Gunnar Nielsen Aaby    M   24     NaN     NaN         Denmark  DEN   
4   4  Edgar Lindenau Aabye    M   34     NaN     NaN  Denmark/Sweden  DEN   

            8     9       10         11          12  \
0        Games  Year  Season       City       Sport   
1  1992 Summer  1992  Summer  Barcelona  Basketball   
2  2012 Summer  2012  Summer     London        Judo   
3  1920 Summer  1920  Summer  Antwerpen    Football   
4  1900 Summer  1900  Summer      Paris  Tug-Of-War   

                             13     14  
0                         Event  Medal  
1   Basketball Men's Basketball    NaN  
2  Judo Men's Extra-Lightweight    NaN  
3   
```

### Question

Print and observe the first 10 records of the dataset.

<!-- grade_id: cell-c2348e0b371280d6, points: 0, grade: false, solution: true -->
```python
print(df_olympics.iloc[:10])
```

Output:
```
0                     1    2    3       4       5               6    7   \
0  ID                  Name  Sex  Age  Height  Weight            Team  NOC   
1   1             A Dijiang    M   24     180      80           China  CHN   
2   2              A Lamusi    M   23     170      60           China  CHN   
3   3   Gunnar Nielsen Aaby    M   24     NaN     NaN         Denmark  DEN   
4   4  Edgar Lindenau Aabye    M   34     NaN     NaN  Denmark/Sweden  DEN   

            8     9       10         11          12  \
0        Games  Year  Season       City       Sport   
1  1992 Summer  1992  Summer  Barcelona  Basketball   
2  2012 Summer  2012  Summer     London        Judo   
3  1920 Summer  1920  Summer  Antwerpen    Football   
4  1900 Summer  1900  Summer      Paris  Tug-Of-War   

                             13     14  
0                         Event  Medal  
1   Basketball Men's Basketball    NaN  
2  Judo Men's Extra-Lightweight    NaN  
3       Football Men's Football    NaN 
```

### Question

Print and observe 5 random records of the dataset.

<!-- grade_id: cell-ef82a27734491ed4, points: 0, grade: false, solution: true -->
```python
df_olympics.sample(n=5)
```

Output:
```
0                     1    2    3       4       5               6    7   \
0  ID                  Name  Sex  Age  Height  Weight            Team  NOC   
3   3   Gunnar Nielsen Aaby    M   24     NaN     NaN         Denmark  DEN   
4   4  Edgar Lindenau Aabye    M   34     NaN     NaN  Denmark/Sweden  DEN   
1   1             A Dijiang    M   24     180      80           China  CHN   
2   2              A Lamusi    M   23     170      60           China  CHN   

            8     9       10         11          12  \
0        Games  Year  Season       City       Sport   
3  1920 Summer  1920  Summer  Antwerpen    Football   
4  1900 Summer  1900  Summer      Paris  Tug-Of-War   
1  1992 Summer  1992  Summer  Barcelona  Basketball   
2  2012 Summer  2012  Summer     London        Judo   

                             13     14  
0                         Event  Medal  
3       Football Men's Football    NaN  
4   Tug-Of-War Men's Tug-Of-War   Gold  
1   Basketball Men's Basketball    NaN 
```

```python

```